Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Leon Müller"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    n = len(A)
    if n == 0:   # check that list is not empty
        return("Error. The list is empty.")
    global_opt = 0   # global_opt index at 0, if no local_opt is every greater than 0, no investment should be made
    for i in range(n):  # loop for each element once
        for j in range(n+1): # loop for each element after i
            if A[i:j]: 
                local_opt = sum(A[i:j])  # sum of outcomes (losses and gains)
                if local_opt > global_opt:
                    global_opt = local_opt  # assign greater local_opt to global_opt
                    buy = i  # start value of global_opt
                    sell = j-1  # end value of global_opt
    return (global_opt) # would usually also return start and end value, but this is not useful for PCW Session 7
    raise NotImplementedError()

A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)

```

Output:
```
4
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
No, I think there is no need to run the whole algorithm again, even though with the current input size this would be faster than to code an alternative solution.

The divide and conquer solution to the max subarray problem used a process for the cross-over subarray finder that I think could be used here. This process found the maximum subarray in subarray L and R that included mid element or mid+1 (in the case of R). In this case here, we only need to know if there is a new subarray that *includes* the final element B[-1]. Once we find this, we can compare it to the old one and see which is larger. 

The process of finiding the maximum subarray that includes B[-1] requires a loop that incrementally compares subarrays either to a sentinel or to the old max. This process needs to start at and always include B[-1].


## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False.
If the maximum subarray of array A and the maximum subarray of B are not the same, then the maximum subarray must be bigger and hence B[-1] (in this case 8) must be included. Array A is fully encompassed by array B. 

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    # YOUR CODE HERE
    if len(x) == 0:
        return("Enter a NON-EMPTY list.")
    elif len(x) == 1:
        return x
    global_max = x[-1]
    for i in range(len(x)+1,0,-1):  # move back incrementally from the last element of the array
            local_max = sum(x[i:-1])  # create a local variable that can is assingned the size of a subarray for later testing
            if local_max > global_max:  # compare subarray size to global max
                global_max = local_max  # if local is greater than global, global is unpdated
    return global_max

# C = [2,2,4]
# mC = bruteforce_max_subarray(C)
# print((C,mC))
A = [-2, -3, 4, -1, -2, 1]
B = [-2, -3, 4, -1, -2, 1, 8]

incremental_max_subarray(B, bruteforce_max_subarray(A))
#     raise NotImplementedError()
```

Output:
```
8
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    # YOUR CODE HERE
    
    raise NotImplementedError()
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
With large inputs, the divide an conquer approach would definitily be more efficient. 